<problem>
  <text>
    <script type="loncapa/python">

# random variables (no need to import random library)
p_men = random.randrange(6,8,1)
p_women = random.randrange(3,5,1)
p_male_given_colorblind = p_men/(p_men+p_women)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1. 
solution1 = "{0}/({0}+{1})".format(p_men, p_women)


# Group all solutions into a list
solutions = [solution1]


# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Suppose that there are equal numbers of men and women in the world, and that \($p_men\)% of men are colorblind \($p_women\)% of women are colorblind. A person is chosen at random and found to be colorblind. What is the probability that the person is male? </p>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>Reverse the condition:\[\textbf{Pr}(\text{male}|\text{colorblind}) = \frac{Pr(\text{male}, \text{colorblind)}}{\textbf{Pr}(\text{colorblin)}} = \]\[\frac{\textbf{Pr}(\text{colorblind}|\text{male})\textbf{Pr}(\text{male})}{\textbf{Pr}(\text{colorblind})}\]</p><p>Law of total probability + Bayes rule:\[\textbf{Pr}(\text{colorblind}) = \text{Pr}(\text{male}, \text{colorblind}) + \textbf{Pr}(\text{female}, \text{colorblind}) =\]\[\textbf{Pr}(\text{colorblind}|\text{male})\textbf{Pr}(\text{male}) + \textbf{Pr}(\text{colorblind}|\text{female})\textbf{Pr}(\text{female})\]</p><hr />
    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>