<problem>
  <text>
    <script type="loncapa/python">

# random variables (no need to import random library)
r1 = random.randrange(10,20,1)
r2 = random.randrange(5,r1,1)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1.
solution1 = "{0}".format(r1)
solution2 = "{0}".format(r2)
solution3 = "{0}/{1}".format(r2,r1)

# Group all solutions into a list
solutions = [solution1, solution2, solution3]

# hint id
hint_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;
      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];
      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
      //fill in hint text id. Need to match the hint_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Suppose we have $r1 cards, numbered 1,...,$r1. Suppose that when we pick a card at random the distribution of the card number is uniform. Let \(A\) be the event that number of the card that we picked is smaller or equal to $r2. What is the probability of the event \(A\)?</p><ul><li><strong>Copy</strong> the number of cards here</li></ul>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li><strong>Copy</strong> the maximal card value in \(A\) here</li></ul>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>The probability of the event is equal to?</li></ul>

    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint import cluster_functions
from hint import universal_hints as u_hint
from hint import conditional_hints as c_hint
from hint import format as string_format

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  param = cluster_functions.make_params(solutions[i], ans)
  if param == {}:
    return
  hint = u_hint.universal_hint(param)
  if hint:
    hint = string_format.format_u_hint(hint)
  else:
    hint_text, hint_answer = c_hint.conditional_hint(param)
    if not hint_text:
      return
    hint = string_format.format_c_hint(hint, hint_id, hint_answer, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
    </script>

  </text>
</problem>